## Exercise 05.1 (random numbers)

- Using the '`randint`' function from the '`random`' module (https://docs.python.org/3/library/random.html#random.randint), 
  develop a function `dice_roll` that emulates the roll of a dice with $n$ sides.

- For $n=6$, devise and implement a test to check that it is a fair dice.

#### (a) Dice roll code:

<!-- grade_id: cell-70a83466894475ee, points: 0, grade: false, solution: true -->
```python
import random

def dice_roll(x):
    return random.randint(1, x)
```

<!-- grade_id: cell-0970743b0c035f78, points: 0, grade: true, solution: false -->
```python
for n in range(1, 20):
    for j in range(100):
        value = dice_roll(n) 
        assert value >= 1 and value <= n
```

#### (b) Test for fairness

<!-- grade_id: cell-0826760c62c21bff, points: 0, grade: true, solution: true -->
```python
dice_result = [0, 0, 0, 0, 0, 0]
n = 6
for i in range(100000):
    roll = dice_roll(n)
    dice_result[roll-1] += 1
print(dice_result)

```

Output:
```
[16814, 16492, 16681, 16591, 16730, 16692]
```

## Exercise 05.2 (data compression)

For devices with limited memory, data compression can be important. Data compression is
a field of its own, but with libraries we can compress (and uncompress) data easily without being expert in
the details.

Below is a program code for compressing a passage from Hamlet, by Shakespeare.

```python
# Import the compression module
import zlib

# Create a string that we wish to compress
text = """
Welcome, dear Rosencrantz and Guildenstern!
Moreover that we much did long to see you,
The need we have to use you did provoke
Our hasty sending. Something have you heard
Of Hamlet's transformation; so call it,
Sith nor the exterior nor the inward man
Resembles that it was. What it should be,
More than his father's death, that thus hath put him
So much from the understanding of himself,
I cannot dream of: I entreat you both,
That, being of so young days brought up with him,
And sith so neighbour'd to his youth and havior,
That you vouchsafe your rest here in our court
Some little time: so by your companies
To draw him on to pleasures, and to gather,
So much as from occasion you may glean,
Whether aught, to us unknown, afflicts him thus,
That, open'd, lies within our remedy."""

# Convert Python string to bytes and check type
text_bytes = text.encode("utf-8")
print(type(text_bytes))

# Get number of bytes used to store string
print("Number of bytes for uncompressed string:", len(text_bytes))

# Compress string and get number of byes used for compressed string
text_comp = zlib.compress(text_bytes)
print("Number of bytes for compressed string:", len(text_comp))

# Display the compression efficiency
print("Compression efficiency: ", len(text_comp)/len(text_bytes))

# Decompress the string
text_decomp = zlib.decompress(text_comp)

# Check that original and decompressed string are the same (more on aseret)
if text != text_decomp.decode("utf-8"):
    print("Problem: original and decompressed string differ.")
```

Using the above as a guide, examine the compression efficiency of 

1. Compressing one large string made up of the the passage by Shakespeare repeated 100 times; and
2. Compressing a random string of the same length as the repeated Shakespeare passage.

To help you, the below function generates random string of length `N`:

```python
import random
import string

def random_string(N):
    return ''.join([random.choice(string.ascii_letters + string.digits) for n in range(N)])

print(random_string(8))
```

### Solution

```python
# Create a string
text = """
Welcome, dear Rosencrantz and Guildenstern!
Moreover that we much did long to see you,
The need we have to use you did provoke
Our hasty sending. Something have you heard
Of Hamlet's transformation; so call it,
Sith nor the exterior nor the inward man
Resembles that it was. What it should be,
More than his father's death, that thus hath put him
So much from the understanding of himself,
I cannot dream of: I entreat you both,
That, being of so young days brought up with him,
And sith so neighbour'd to his youth and havior,
That you vouchsafe your rest here in our court
Some little time: so by your companies
To draw him on to pleasures, and to gather,
So much as from occasion you may glean,
Whether aught, to us unknown, afflicts him thus,
That, open'd, lies within our remedy."""
```

Import the necessary modules:

```python
import random
import string
import zlib
```

Repeat the Shakespeare string 100 times, and compress:

<!-- grade_id: cell-c08bb020d9be6620, points: 0, grade: false, solution: true -->
```python
import random
import string
import zlib

# Create string of Shakespeare passage repeated 100 times
text = """
Welcome, dear Rosencrantz and Guildenstern!
Moreover that we much did long to see you,
The need we have to use you did provoke
Our hasty sending. Something have you heard
Of Hamlet's transformation; so call it,
Sith nor the exterior nor the inward man
Resembles that it was. What it should be,
More than his father's death, that thus hath put him
So much from the understanding of himself,
I cannot dream of: I entreat you both,
That, being of so young days brought up with him,
And sith so neighbour'd to his youth and havior,
That you vouchsafe your rest here in our court
Some little time: so by your companies
To draw him on to pleasures, and to gather,
So much as from occasion you may glean,
Whether aught, to us unknown, afflicts him thus,
That, open'd, lies within our remedy."""

text *= 100

text_bytes = text.encode("utf-8")
text_comp = zlib.compress(text_bytes)
```

Output:
```
0.01178343949044586
```

Create random string and compress:

<!-- grade_id: cell-c5d8c07589ec3cd1, points: 0, grade: false, solution: true -->
```python
import random
import string
import zlib

def random_string(N):
    return ''.join([random.choice(string.ascii_letters + string.digits) for n in range(N)])

string_bytes = random_string(len(text)).encode("utf-8")
string_comp = zlib.compress(string_bytes)
```

Compare compression efficiency:

<!-- grade_id: cell-a2dd98ee370e22a3, points: 0, grade: true, solution: true -->
```python
print(len(text_comp)/len(text_bytes))
print(len(string_comp)/len(string_bytes))
```

Output:
```
0.01178343949044586
0.7520891719745223
```