## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
# def area(x0, y0, x1, y1, x2, y2):
#     return abs((x0*(y1-y2)+x1*(y2-y0)+x2*(y0-y1))/2)

# a = [x0, y0]
# b = [x1, y1]
# c = [x2, y2]

vert1 = [0, 0]
vert2 = [0, 2]
vert3 = [3, 0]

def area(a, b, c):
    if len(a) != 2 and len(b) == 2 and len(c) == 2:
        return "Please ensure vert1 has exactly 2 coordinates, x and y."
    elif len(a) == 2 and len(b) != 2 and len(c) == 2:
        return "Please ensure vert2 has exactly 2 coordinates, x and y."
    elif len(a) == 2 and len(b) == 2 and len(c) != 2:
        return "Please ensure vert3 has exactly 2 coordinates, x and y."
    else:
        return abs((a[0]*(b[1]-c[1])+b[0]*(c[1]-a[1])+c[0]*(a[1]-b[1]))/2)

print(area(vert1, vert2, vert3))
```

Output:
```
3.0
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
#vert = [[x0, y0],
#        [x1, y1],
#        [x2, y2],
#        [x3, y3]]
vert = [[1, 1],
        [1, 2],
        [2, 2],
        [2, 1]]

def area(n):
    A = 0
    if n == len(vert):
        for i in range(n):
            #i=0 to i=n-1
            if i != n-1:
                A += (vert[i][0])*(vert[i+1][1])-(vert[i+1][0])*(vert[i][1])
            else:
                A += (vert[i][0])*(vert[0][1])-(vert[0][0])*(vert[i][1])
        return 0.5*abs(A)
    elif n > len(vert):
        print("Please provide sufficient number of vertices")
    else:
        print("The number of vertices provided is more than the number you have specified")


area(4)
```

Output:
```
1.0
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    "Return sum of two vectors"
    new_vector = []
    if len(x) == len(y):
        for i in range(len(x)):
          new_vector.append((x[i]+y[i]))
        return new_vector
    else:
        return "size mismatch"
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

```python
# a = [1, 2, 3]
# b = [4, 6, 8]
# c = [5, 8, 11]

c = []
a = "test"
c.append(a)

def sum_vector(a, b):
  c = [1, 3, 5, 23, 4 6, 7, 8 ...... 98]
  for i in 
  c[0]=a[0]+b[0]

  return c

```

### Extension: list comprehension

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A dictionary from college abbreviation to name; and
1. A list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
name_abb = {"Emmanuel":"EM",
            "Chuchill":"CHU",
            "Clare":"CL",
            "Darwin":"DAR",
            "Downing":"DOW"}
print(name_abb)

abb_name = dict((i, j) for j, i in name_abb.items())
print(abb_name)
  
list = []
for x in name_abb:
  list.append(name_abb[x])
  list.sort()
print(list)
```

Output:
```
{'Emmanuel': 'EM', 'Chuchill': 'CHU', 'Clare': 'CL', 'Darwin': 'DAR', 'Downing': 'DOW'}
{'EM': 'Emmanuel', 'CHU': 'Chuchill', 'CL': 'Clare', 'DAR': 'Darwin', 'DOW': 'Downing'}
['CHU', 'CL', 'DAR', 'DOW', 'EM']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```