## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n): 
    "Compute the nth Fibonacci number using recursion"
    try:
      if n < 0:
        raise ValueError("n has to be >= 0")
      elif n == 0:
        return 0  # This doesn't call f, so it breaks out of the recursion loop
      elif n == 1:
        return 1  # This doesn't call f, so it breaks out of the recursion loop
      else:
        return f(n - 1) + f(n - 2)  # This calls f for n-1 and n-2 (recursion), and returns the sum
    except:
      return "The Fibonacci series is valid only for n ≥ 0. Please try again"
      
print(f(10))
```

Output:
```
55
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
# Perform some tests    
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

Output:
```
/usr/local/lib/python3.7/dist-packages/_pytest/python_api.py in __exit__(self, *tp)
    665         __tracebackhide__ = True
    666         if tp[0] is None:
--> 667             fail(self.message)
    668         self.excinfo.__init__(tp)
    669         suppress_exception = issubclass(self.excinfo.type, self.expected_exception)

/usr/local/lib/python3.7/dist-packages/_pytest/outcomes.py in fail(msg, pytrace)
     95     """
     96     __tracebackhide__ = True
---> 97     raise Failed(msg=msg, pytrace=pytrace)
     98 
     99 

Failed: DID NOT RAISE <class 'ValueError'>
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    "Evaluate polynomial function"
    return x**3 - 6*x**2 + 4*x + 12

def compute_root(f, x0, x1, tol, max_it):

    try:
      "Compute roots of a function using bisection"
      error = tol + 1
      it = 0
      while error > tol:
          it += 1
          x_mid = (x0 + x1)/2
          #print("x_mid is", x_mid)
          f0 = my_f(x0)
          f_mid = my_f(x_mid)
          #print("f_mid is", f_mid)

          if f0*f_mid < 0:
            x1 = x_mid
          else:
            x0 = x_mid

          if abs(f_mid) < tol:
            break
          else:
            pass

          if it > max_it:
            raise ValueError("max iterations reached")
      return x_mid, f_mid, it

    except:
      return "Oops, iteration count is very large. Breaking out of while loop."

print(compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=30))
```

Output:
```
(4.534070134162903, -7.047073751209609e-07, 23)
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=20)
```