## Exercise 12.1

Create a class to represent vectors of arbitrary length and which is initialised with a list of values, e.g.:
```python
x = MyVector([0, 2, 4])
```

Equip the class with methods that:

1. Return the length of the vector (use name `size`)
2. Compute the norm of the vector $\sqrt{x \cdot x}$ (use name `norm`)
3. Compute the dot product of the vector with another vector (use name `dot`)

Test your implementation using two vectors of length 3. To help you get started, a skeleton of the class is provided below. Don't forget to use `self` where necessary.

```python
class MyVector:
    """A vector object that can return its size and norm, and can compute the dot product 
    with another vector  """
    
    def __init__(self, x):
        self.x = x
        
    # Return length of vector
    def size(self):
        # Add your code here
        pass  # This can be removed once the body is added
    
    # This allows access by index, e.g. y[2]
    def __getitem__(self, index):
        return self.x[index]

    # Return norm of vector
    def norm(self):
        # Add your code here
        pass  # This can be removed once the body is added
    
    # Return dot product of vector with another vector
    def dot(self, other):
        # Add your code here
        pass  # This can be removed once the body is added
```

### Solution

<!-- grade_id: cell-ff4d6cba19bb1b8e, points: 0, grade: false, solution: true -->
```python
import math  # Use this to get the sqrt
import numpy as np

class MyVector:
    """A vector object that can return its size and norm, and can compute the dot product 
    with another vector  """

    def __init__(self, x):
        self.x = x

    def size(self):
        return len(self.x)

    def __getitem__(self, index):
        return self.x[index]
    
    def norm(self):
        return math.sqrt(sum(i**2 for i in self.x))     

    def dot(self, other):
        return sum(self.x[i]*other[i] for i in range(len(self.x)))
```

<!-- grade_id: cell-b601b5cc8cf8ab44, points: 0, grade: true, solution: false -->
```python
# Create two vectors
u = MyVector([1, 1, 2])
v = MyVector([2, 1, 1])

assert u.size() == 3
assert round(u.norm() - 2.449489742783178) == 0.0
assert round(u.dot(v) - 5.0, 10) == 0.0 
```

## Exercise 12.2

1. Create a class for holding a student record entry. It should have the following attributes:
   - Surname
   - Forename
   - Birth year
   - Tripos year
   - College
   - CRSid (optional field)
1. Equip your class with the method '`age`' that returns the age of the student in whole years
1. Equip your class with the method '`__repr__`' such using `print` on a student record displays with the format

       Surname: Bloggs, Forename: Andrea, College: Churchill

1. Equip your class with the method `__lt__(self, other)` so that a list of record entries can be sorted by 
   (surname, forename) using the Python built-in sort function. 
   
   Create a list of entries and test the sorting. Make sure you have two entries with the same
   surname.

> Recall that the methods starting with `__`, e.g. `__lt__` and `__repr__`, should **not** be called directly. 
> Python with map them to other operations, e.g. `__repr__` is called when using `print`, and `__lt__` is called 
> when using `<`. These functions must have a return value.

*Hint:* To get the current year:

```python
import datetime
year = datetime.date.today().year
print(year)
```

Output:
```
2021
```

### Solution

<!-- grade_id: cell-6e857f6089fe0ba8, points: 0, grade: false, solution: true -->
```python
import datetime 
class StudentEntry:
    def __init__(self, surname, forename, birth_year, year, college, crsid=None):
        self.surname = surname
        self.forename = forename
        self.birth_year = birth_year
        self.year = year
        self.college = college
        self.crsid = crsid

    def age(self):
        return datetime.date.today().year - self.birth_year

    def __repr__(self):
        return "Surname: " + str(self.surname) + ", " + "Forename: " + str(self.forename) + ", " + "College: " + str(self.college)

    # Custom '<' operator (method)
    def __lt__(self, other):
        if self.surname != other.surname:
            return self.surname < other.surname
        else:
            return self.forename < other.forename

s0 = StudentEntry("Bloggs", "Andrea", 1996, 1, "Churchill", "ab1001")       
s1 = StudentEntry("Reali", "John", 1997, 1, "Corpus Christi")       
s2 = StudentEntry("Bacon", "Kevin", 1996, 1, "Newnham")
s3 = StudentEntry("Bacon", "Alexander", 1996, 1, "Queens")

s = [s0, s1, s2, s3]
s.sort()
print(s)
```

Output:
```
[Surname: Bacon, Forename: Alexander, College: Queens, Surname: Bacon, Forename: Kevin, College: Newnham, Surname: Bloggs, Forename: Andrea, College: Churchill, Surname: Reali, Forename: John, College: Corpus Christi]
```

<!-- grade_id: cell-ce8aafb3e627dd86, points: 0, grade: true, solution: false -->
```python
s0 = StudentEntry("Bloggs", "Andrea", 1996, 1, "Churchill", "ab1001")       
s1 = StudentEntry("Reali", "John", 1997, 1, "Corpus Christi")       
s2 = StudentEntry("Bacon", "Kevin", 1996, 1, "Newnham")
s3 = StudentEntry("Bacon", "Alexander", 1996, 1, "Queens")

assert s0 < s1
assert s0 > s2
assert s3 < s2
assert s0.age() ==  datetime.date.today().year - 1996 
assert s1.age() ==  datetime.date.today().year - 1997
assert str(s1) == "Surname: Reali, Forename: John, College: Corpus Christi"

# Test sorting
s = [s0, s1, s2, s3]
s.sort()
for earlier, later in zip(s, s[1:]):
    assert earlier.surname <= later.surname
    if earlier.surname == later.surname:
        assert earlier.forename <= later.forename
```